---
  title: "негиздерин Intro"
  description: "This is a template chapter."
---

## Бул кантип иштейт

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 206ec04c4b   
```


укугу боюнча Сиз редактордо көнүгүүлөрдү чечүүгө R кодун терип керек. сен да: "Жооп тапшыруу 'баскычын тийгенде, кодексинин ар бир сабы чечмеленет жана R менен жол-жоболоштурулат жана коду туура болду бекен же жокпу, бир кабар алуу. Сиздин R кодексин иштеп чыгуу төмөнкү оң бурчунда консолго көрсөтүлгөн.

R сен жана башкалар R коду эмне жөнүндө болуп жатканын түшүнө алат деп, комментарий кошуу үчүн # белгисин колдонууга болот. Эле Twitter сыяктуу! Comments R код катары эмес, ошондуктан алар натыйжага таасир тийгизе албайт. Мисалы, укук боюнча редактору 4 + 3 комментарийленген эсептөө.

Ошондой эле R түз консолго буйруктарды аткарууга болот. Бул сиздин берүүнүн тууралыгы текшерилет эмес, ошондой эле R кодексине, татып жакшы жолу болуп саналат.


`@instructions`
- укугу боюнча редактору буга чейин кээ бир үлгү коду бар. Сиз саптары чыныгы R коду жана талкуулар болуп саналган кайсы көрө алабы?
- 6 жана 12 суммасын эсептейт кодексинин бир сапты кошуп, "Жооп тапшыруу" баскычты бас.

`@hint`
- Жөн гана берилген кодду, мисалы болуп, 6, 12 суммасын эсептейт R кодексинин бир сапты кошуп!

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```

`@sample_code`

```{r}
# Эсептеп 3 + 4
3 + 4

# Эсептеп 6 + 12
```

`@solution`

```{r}
# Эсептеп 3 + 4
3 + 4

# Эсептеп 6 + 12
6 + 12
```

`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
