---
  title: "негиздерин Intro"
  description: "This is a template chapter."
---

## Бул кантип иштейт

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 206ec04c4b   
```


укугу боюнча Сиз редактордо көнүгүүлөрдү чечүүгө R кодун терип керек. сен да: "Жооп тапшыруу 'баскычын тийгенде, кодексинин ар бир сабы чечмеленет жана R менен жол-жоболоштурулат жана коду туура болду бекен же жокпу, бир кабар алуу. Сиздин R кодексин иштеп чыгуу төмөнкү оң бурчунда консолго көрсөтүлгөн.

R сен жана башкалар R коду эмне жөнүндө болуп жатканын түшүнө алат деп, комментарий кошуу үчүн # белгисин колдонууга болот. Эле Twitter сыяктуу! Comments R код катары эмес, ошондуктан алар натыйжага таасир тийгизе албайт. Мисалы, укук боюнча редактору 4 + 3 комментарийленген эсептөө.

Ошондой эле R түз консолго буйруктарды аткарууга болот. Бул сиздин берүүнүн тууралыгы текшерилет эмес, ошондой эле R кодексине, татып жакшы жолу болуп саналат.


`@instructions`
- укугу боюнча редактору буга чейин кээ бир үлгү коду бар. Сиз саптары чыныгы R коду жана талкуулар болуп саналган кайсы көрө алабы?
- 6 жана 12 суммасын эсептейт кодексинин бир сапты кошуп, "Жооп тапшыруу" баскычты бас.

`@hint`
- Жөн гана берилген кодду, мисалы болуп, 6, 12 суммасын эсептейт R кодексинин бир сапты кошуп!

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```

`@sample_code`

```{r}
# Эсептеп 3 + 4
3 + 4

# Эсептеп 6 + 12
```

`@solution`

```{r}
# Эсептеп 3 + 4
3 + 4

# Эсептеп 6 + 12
6 + 12
```

`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Р эсеп

```yaml
type: NormalExercise 
xp: 100 
key: 216dbc50fb   
```


анын негизги түрүндө R жөнөкөй эсептегич катары колдонсо болот. Төмөнкү эсеп операторлорду карап көрөлү:

Кошулган: +
Subtraction: -
Multiplication: *
Бөлүм: /
Даражага: ^
Modulo: %%
Акыркы эки айрым түшүндүрүү керек:

^ Оператору өзүнүн укугун санынын бийликке солго жаралат: мисалы, 3 ^ 2 9.
модулдук үлгү 5 Modulo үчүн, анын укугу катар сол санына бөлүү калган кайтарып 3 же 5 %% 3 2.
Бул билим менен, көнүгүүнү аяктоо үчүн төмөнкү көрсөтмөлөрдү аткарышат.


`@instructions`
- бийликке 2 эсептеп редактордо ^ 5 2, 5 тиби.
- 28 модулдук эсептөө терип 28 %% 6 6.
- "Жооп тапшыруу" Click жана консолго R өндүрүшүнүн карап бар.
- # Белгиси R кодексинин боюнча комментарий кошуу үчүн колдонулат карап көрөлү.

`@hint`
модулдук операторунун дагы бир мисал: 9 %% 2 барабар 1.

`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}
# Кошумча
5 + 5

# A кемитүү
5 - 5

# A көбөйтүү
3 * 5

  # Пикир келишпестик
(5 + 5) / 2

# Exponentiation


# Modulo

```

`@solution`

```{r}
# Кошумча
5 + 5

# A кемитүү
5 - 5

# A көбөйтүү
3 * 5

  # Пикир келишпестик
(5 + 5) / 2

# Exponentiation
2 ^ 5

# Modulo
28 %% 6
```

`@sct`

```{r}

```

